## Exercise 02.1 (if-else)

Consider the following assessment criteria which map a score out of 100 to an 
assessment grade:

| Grade            | Raw score  (/100)      |
| ---------------- | ---------------------- |
| Excellent        | $\ge 85$               |
| Very good        | $\ge 76.5$ and $< 85$  |
| Good             | $\ge 64$ and $< 76.5$  |
| Need improvement | $\ge 40$ and $< 64$    |
| Did you try?     | $< 40$                 |

Write a program that, given an a score, prints the appropriate grade. Print an error message if the input score is greater than 100 or less than zero.

<!-- grade_id: cell-71121f470be22e4a, points: 0, grade: false, solution: true -->
```python
# Score from user
score = 72
## defining a function that checks a score inputted into it ##
def check(score):
    ## checks if score inputted is greater than 100 or less than zero ##
    if score < 0 or score > 100:
        print("Invalid input")
    ## checks if lower than 40 ##
    elif score < 40:
        print("Did you try?")
    ## checks if lower than 64 ##
    elif score < 64:
        print("Need improvement")
    ## checks if lower than 76.5 ##
    elif score < 76.5:
        print("Good")
    ## checks if lower than 85 ##
    elif score < 85:
        print("Very good")
    ## otherwise score is higher than 85 ##
    else:
        print("Excellent")

## using previously defined function to check function works ##
check(score)
check(-1)
check(103)
check(32)
check(50)
check(70)
check(80)
check(90)
```

Output:
```
Good
Invalid input
Invalid input
Did you try?
Need improvement
Good
Very good
Excellent
```

## Exercise 02.2 (bisection)

Bisection is an iterative method for finding approximate roots of a function. Say we know that the function $f(x)$ has one root between $x_{0}$ and $x_{1}$ ($x_{0} < x_{1}$). We then:

- Evaluate $f$ at the midpoint $x_{\rm mid} = (x_0 + x_1)/2$, i.e. compute
   $f_{\rm mid} = f(x_{\rm mid})$
- Evaluate $f(x_0) \cdot f(x_{\rm mid})$

  - if $f(x_0) \cdot f(x_{\rm mid}) < 0$: 

    $f$ must change sign somewhere between $x_0$ and $x_{\rm mid}$, hence the root must lie between 
    $x_0$ and $x_{\rm mid}$, so set $x_1 = x_{\rm mid}$.
   
  - else:

    $f$ must change sign somewhere between $x_{\rm mid}$ and $x_1$, so set
    $x_0 = x_{\rm mid}$.

The above steps can be repeated a specified number of times, or until $|f_{\rm mid}|$
is below a tolerance, with $x_{\rm mid}$ being the approximate root.


### Task

The function

$$
f(x) = \frac{x^{5}}{10} + x^3 - 10x^2 + 4x + 7
$$


has one root in the range $0 < x < 2$.

1. Use the bisection method to find an approximate root $x_{r}$ using 20 iterations 
   (use a `for` loop).
2. Use the bisection method to find an approximate root $x_{r}$ such that 
   $\left| f(x_{r}) \right| < 1 \times 10^{-6}$ and report the number of iterations 
   required (use a `while` loop).

Store the approximate root using the variable `x_mid`, and store $f(x_{\rm mid})$ using the variable `f`.

*Hint:* Use  `abs` to compute the absolute value of a number, e.g. `y = abs(x)` assigns the absolute value of `x` to `y`. 

#### (1) Using a `for` loop.

<!-- grade_id: cell-3a0d0a59613d6b3f, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

# Use 20 iterations
for n in range(20):
    # Compute midpoint
    x_mid = (x0 + x1) / 2

    # Evaluate function at (i) left end-point and at (ii) midpoint
    f0 = (x0**5) / 10 + x0**3 - 10 * x0**2 + 4 * x0 + 7
    f = (x_mid**5) / 10 + x_mid**3 - 10 * x_mid**2 + 4 * x_mid + 7

    # Evaluating whether f0 * f < 0
    if f0*f < 0:
        x1 = x_mid
    else:
        x0 = x_mid

    print(n, x_mid, f)
```

Output:
```
0 1.0 2.0999999999999996
1 1.5 -5.365625000000001
2 1.25 -1.36669921875
3 1.125 0.4477813720703132
4 1.1875 -0.4408627510070797
5 1.15625 0.008327355980872753
6 1.171875 -0.21507753478363156
7 1.1640625 -0.1030741602036862
8 1.16015625 -0.04729774910292761
9 1.158203125 -0.019466230897839054
10 1.1572265625 -0.005564689502373099
11 1.15673828125 0.0013825210451265946
12 1.156982421875 -0.0020907873792292975
13 1.1568603515625 -0.00035405894194262544
14 1.15679931640625 0.0005142496094645566
15 1.156829833984375 8.009997302949046e-05
16 1.1568450927734375 -0.00013697832466341708
17 1.1568374633789062 -2.8438885864900953e-05
18 1.1568336486816406 2.5830616070976475e-05
19 1.1568355560302734 -1.304116775457942e-06
```

<!-- grade_id: cell-421824bbd4e61fae, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
assert math.isclose(x_mid, 1.1568355560302734)
assert abs(f) < 1e-5
```

#### (2) Using a `while` loop

Use the variable `counter` for the iteration number. 

*Remember to guard against infinite loops.*

<!-- grade_id: cell-1747a3fb1dd1de48, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

tol = 1.0e-6
error = tol + 1.0

# initialising the f(x) defined
def fx(x):
    answer = (x**5) / 10 + x**3 - 10 * x**2 + 4 * x + 7
    return answer

# Iterate until tolerance is met
counter = 0
while error > tol:
    # Compute midpoint
    x_mid = (x0 + x1) / 2

    # Evaluate function at (i) left end-point and at (ii) midpoint
    f0 = fx(x0)
    f = fx(x_mid)

    # same checks as last time
    if f0*f < 0:
        x1 = x_mid
    else:
        x0 = x_mid

    # adding one to counter
    counter += 1

    #seeing if f(x_mid) < 1e-6
    if abs(f) < 1e-6:
        print("Breaking")
        print(counter, x_mid, error)
        break
    
    # Guard against an infinite loop
    if counter > 1000:
        print("Oops, iteration count is very large. Breaking out of while loop.")
        break
    
    print(counter, x_mid, error)
```

Output:
```
1 1.0 1.000001
2 1.5 1.000001
3 1.25 1.000001
4 1.125 1.000001
5 1.1875 1.000001
6 1.15625 1.000001
7 1.171875 1.000001
8 1.1640625 1.000001
9 1.16015625 1.000001
10 1.158203125 1.000001
11 1.1572265625 1.000001
12 1.15673828125 1.000001
13 1.156982421875 1.000001
14 1.1568603515625 1.000001
15 1.15679931640625 1.000001
16 1.156829833984375 1.000001
17 1.1568450927734375 1.000001
18 1.1568374633789062 1.000001
19 1.1568336486816406 1.000001
20 1.1568355560302734 1.000001
21 1.156834602355957 1.000001
22 1.1568350791931152 1.000001
23 1.1568353176116943 1.000001
Breaking
24 1.1568354368209839 1.000001
```

<!-- grade_id: cell-ddff32f44df88388, points: 0, grade: true, solution: false -->
```python
## tests ##
assert counter == 24
assert abs(f) < 1.0e-6
```

## Exercise 02.3 (series expansion)

For $|x| < 1$ the series: 

$$
(1 + x)^{-1/2} = \sum_{n = 0}^{\infty} \frac{(-1)^n (2n)!}{4^n (n!)^2} x^n
$$

converges.

1. Using a `for` statement, approximate $1/\sqrt{0.16}$ using 30 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $1/\sqrt{0.16}$ to within $1 \times 10^{-5}$. 

Store the absolute value of the error in the variable `error`.

### Hints

To compute the factorial, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
```

You only need `import math` once at the top of your program. Standard modules, like `math`, will be explained in a later

<!-- The power series expansion for the sine function is: 

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

(See mathematics data book for a less compact version; this compact version is preferred here as it is simpler to program.)

1. Using a `for` statement, approximate $\sin(3\pi/2)$ using 15 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $\sin(3\pi/2)$ to within $1 \times 10^{-8}$. 

Store the absolute value of the error in the variable `error`.

*Note:* Calculators and computers use iterative or series expansions to compute trigonometric functions, similar to the one above (although they use more efficient formulations than the above series).

### Hints

To compute the factorial and to get a good approximation of $\pi$, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
pi = math.pi
```
You only need '`import math`' once at the top of your program. Standard modules, like `math`, will be explained in a later. If you want to test for angles for which sine is not simple, you can use 
```python
a = 1.3
s = math.sin(a)
```    
to get an accurate computation of sine to check the error. -->

#### (1) Using a `for` loop

<!-- grade_id: cell-d6df4b428f969a5d, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.factorial
import math

# Value of x (such that (1 + x) = 0.16  
x = -0.84

# Actual answer
actual = 2.5

# Initialise approximation of the function
approx_f = 0.0

# Function defining the approximation
def approximation(n):
    answer = ((-1)**n * math.factorial(2*n)) /(4**n * math.factorial(n)**2) * x**n
    return answer

# looping 30 iterations
for i in range(30):
    approx_f += approximation(i)
    print(approx_f)

# calculating error 
error = actual-approx_f
    
# printing error
print(f"The error is: {error}")
```

Output:
```
1.0
1.42
1.6845999999999999
1.8698199999999998
2.0059566999999996
2.1088760451999997
2.1881239410039997
2.2499372997311196
2.2986153197287265
2.3372332155934945
2.3680502964935797
2.3927599922698297
2.412651297369711
2.428717351488846
2.441730855325345
2.4522978204405823
2.460896688303107
2.4679072946898475
2.473632623239019
2.478315339326078
2.4821504838013797
2.485295302271127
2.4878769123422013
2.4899983223571276
2.4917431820944045
2.493179550630131
2.4943628973237946
2.495338500931237
2.4961433739073775
2.4968078104263016
The error is: 0.0031921895736983785
```

<!-- grade_id: cell-8b269b717e92414d, points: 0, grade: true, solution: false -->
```python
## test ##
assert error >= 0
assert error < 1.0e-2
```

#### (2) Using a `while` loop

*Remember to guard against infinite loops.*

<!-- grade_id: cell-03df25d83df3674e, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.sin and math.factorial
import math

# Value of x (such that (1 - x) = 0.16)
x = -0.84

# Actual answer
actual = 2.5

# Tolerance and initial error (this just needs to be larger than tol)
tol = 1.0e-5
error = tol + 1.0

# Initialise approximation of function
approx_f = 0.0

# Function defining the approximation
def approximation(n):
    answer = ((-1)**n * math.factorial(2*n)) /(4**n * math.factorial(n)**2) * x**n
    return answer

# Initialise counter
n = 0

# Loop until error satisfies tolerance, with a check to avoid 
# an infinite loop
while error > tol and n < 1000:
    approx_f += approximation(n)
    print(approx_f)
    error = actual - approx_f
    
    # Increment counter
    n += 1    
    
    
print("\nThe error is:", error)
print("Number of terms in series:", n)
```

Output:
```
1.0
1.42
1.6845999999999999
1.8698199999999998
2.0059566999999996
2.1088760451999997
2.1881239410039997
2.2499372997311196
2.2986153197287265
2.3372332155934945
2.3680502964935797
2.3927599922698297
2.412651297369711
2.428717351488846
2.441730855325345
2.4522978204405823
2.460896688303107
2.4679072946898475
2.473632623239019
2.478315339326078
2.4821504838013797
2.485295302271127
2.4878769123422013
2.4899983223571276
2.4917431820944045
2.493179550630131
2.4943628973237946
2.495338500931237
2.4961433739073775
2.4968078104263016
2.497356634990933
2.4978102119375736
2.4981852633753268
2.4984955332011043
2.498752327109851
2.498964952466293
2.499141077136546
2.499287022606545
2.499408003719833
2.4995083249814676
2.4995915414679937
2.499660590855106
2.4997179018464095
2.499765483297329
2.499804997529524
2.4998378206850673
2.499865092446043
2.4998877570201645
2.499906596947403
2.4999222610011924
2.499935287228324
2.4999461219843027
2.4999551356678342
2.4999626357328104
2.4999688774535516
2.499
```

<!-- grade_id: cell-03df25d83df3674dfwe, points: 0, grade: true, solution: false -->
```python
## test ##
assert error >= 0
assert error <= 1.0e-5
```